# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 26.1 - When a Python program opens a file to read from it, no other programs can read from the same file at the same time.

> True or False


```python
answer261 = False
```

### 26.2 - What is the advantage of opening files in Python using a `with` block?

> A) It ensures the file is closed as soon as you're done using it
>
> B) It runs faster than not using the `with` block
>
> C) It automatically loads the whole file into memory so that it can be accessed faster
>
> D) Actually, it is the only way to do it


```python
answer262 = "A"
```

### 26.3 - Which of the following is true?

> A) Many files only contain text, but many file formats do not use text
>
> B) All files contain only text, it's the just the text contents that explain the data to the computer
>
> C) All files contain only text in either JSON, XML, or CSV.
>
> D) Even if files contain only text, there still has to be a special code at the beginning to explain that this file will contain only text.


```python
answer263 = "A"
```

### 27.1 - Which of the following do you think is least likely to be used as a delimiter in structured files?

> A) `,`
>
> B) `|`
>
> C) `[tab]`
>
> D) `[space]`


```python
answer271 = "D"
```

### 27.2 - Which approaches can be used to read pipe-delimited (`|`) files?

*Report your answer as a list of options, e.g.: `["A","C"]`*

> A) Using the csv module
>
> B) Using the pandas module
>
> C) Using built-in file reading
>
> D) Using the psv module


```python
answer272 = ["A","B","C","D"]
```

### 27.3 - In a CSV file, how does the format change if you want to have a `,` in the actual contents of the file rather than as a delimiter?

> A) You can't do that.  You have to remove every occurence of `,`s
>
> B) Enclose the `,` using quotation marks
>
> C) Escape the `,` using a back-slash
>
> D) Nothing special, it will be obvious to the reader


```python
answer273 = "B"
```

### 28.1 - Select all of the data types that can be represented within a JSON format:

Provide your answer in the form of a list.  For example: `answer = ["A","B"]` 

> A) str
>
> B) list
>
> C) dict
>
> D) int


```python
answer281 = ["A","B","C","D"]
```

### 28.2 - JSON stands for:

> A) Just a Simple Object Notation
>
> B) JavaScript Object Notation
>
> C) JavaScript Objects and Numbers
>
> D) Java Standard Object Notation


```python
answer282 = "B"
```

### 28.3 - Without just running the code, what will the following code print?

```
import json

data = """
[  {"id": { "last name": "Boal", "first name": "Paul", "middle initial": "E" }},
   {"id": { "last name": "Westhus", "first name": "Eric" }} ]
"""

names = json.loads(data)
list(names[1]['id'].keys())[1]
```

> A) 'Eric'
>
> B) 'first name'
>
> C) 'last name'
>
> D) 'Boal'


```python
answer283 = "B"
```

### 29.1 - In this class, XML Stands for:

> A) eXtensible Markup Language
>
> B) The 10th version of the Markup Language
>
> C) eXhange Markup Language
>
> D) X-Men Legends


```python
answer291 = "A"
```

### 29.2 - Which of the following is not a type of data in an XML document?

> A) tags
>
> B) attributes
>
> C) text
>
> D) list


```python
answer292 = "D"
```

## #3 - True or False, each element in an XML document can only have one child with a specific tag.

For example, is the following XML illegal because there are two child elements with the tag "kid":

```xml
<family>
  <parent role="dad">Paul Boal</parent>
  <kid role="daughter">Eleanor</kid>
  <kid role="daughter">Adelaide</kid>
</family>
```

Answer `False` if you think the above XML is valid.  

Answer `True` if you think the above XML is invalid.


```python
answer293 = False
```

### 30.1 List of Allergies

In this directory, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
def allergy_count(json_file):
    with open(json_file) as f:
        allergies = json.load(f)
        
    return(len(allergies.get('entry')))
### END SOLUTION
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count('allergies.json')) == int
### BEGIN HIDDEN TESTS
assert allergy_count('allergies.json') == 4
### END HIDDEN TESTS

```

### 30.2 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
def allergy_list(json_file):
    patients = []
    
    with open(json_file) as f:
        allergies = json.load(f)
    
    entries = allergies.get('entry')
    for entry in entries:
        resource = entry.get('resource')
        name = resource.get('patient').get('display')
        substance = resource.get('substance').get('text')
        reaction = resource.get('reaction')[0].get('manifestation')[0].get('text')
        patients.append([name,substance,reaction])
        
    return(patients)

### END SOLUTION
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list('allergies.json') == output

```

```python
allergy_list('allergies.json')
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

### 30.3 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION
def allergy_reaction(json_file, patient, substance):
    reactions = allergy_list(json_file)
    
    for reaction in reactions:
        if reaction[0].upper() == patient.upper() and reaction[1].upper() == substance.upper():
            return reaction[2]
        
    return None
        
### END SOLUTION
```

```python
allergy_reaction('allergies.json', 'Paul Boal', 'STRAWBERRY')
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction('allergies.json', 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction('allergies.json', 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction('allergies.json', 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction('allergies.json', 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction('allergies.json', 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

```python

```